# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

```python
def Energy():
    #Energy Start
    Energy=pd.read_excel('assets/Energy Indicators.xls')
    Energy = Energy[17:244]
    #Energy.columns
    Energy.columns = Energy.columns[:2].tolist() + ['Country','Energy Supply','Energy Supply per Capita','% Renewable']
    del Energy['Unnamed: 0']
    del Energy['Unnamed: 1']
    #convert
    Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)
    #Nan's
    Energy['Energy Supply'] = Energy['Energy Supply'].replace(r'[.]*', np.NaN, regex=True)
    Energy['% Renewable'] = Energy['% Renewable'].replace(r'[.]*', np.NaN, regex=True)
    Energy['Energy Supply per Capita'] = Energy['Energy Supply per Capita'].replace(r'[.]*', np.NaN, regex=True)
    #Corrections for country
    Energy['Country'] = Energy['Country'].str.replace(r' [(].*[)]',"")
    Energy['Country'] = Energy['Country'].str.replace(r"\d{1,3}","")
    Energy['Country']=Energy['Country'].replace({"Republic of Korea": "South Korea",
                                                 "United States of America": "United States",
                                                 "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                                                 "China, Hong Kong Special Administrative Region": "Hong Kong"})
    #Energy Done
    return Energy
def GDP():
    #GDP Start
    GDP = pd.read_csv('assets/world_bank.csv',
                      header=4)
    #Corrections for country
    GDP['Country Name'] = GDP['Country Name'].replace({"Korea, Rep.": "South Korea",
                                                       "Iran, Islamic Rep.": "Iran",
                                                       "Hong Kong SAR, China": "Hong Kong"})
    GDP.rename(columns={'Country Name':'Country'}, inplace=True)
    #GDP Done
    return GDP
def ScimEn():
    #ScimEn Start
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    #ScimEn Done
    return ScimEn
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    #Energy Start
    Energy=pd.read_excel('assets/Energy Indicators.xls')
    Energy = Energy[17:244]
    #Energy.columns
    Energy.columns = Energy.columns[:2].tolist() + ['Country','Energy Supply','Energy Supply per Capita','% Renewable']
    del Energy['Unnamed: 0']
    del Energy['Unnamed: 1']
    #convert
    Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000)
    #Nan's
    Energy['Energy Supply'] = Energy['Energy Supply'].replace(r'[.]*', np.NaN, regex=True)
    Energy['% Renewable'] = Energy['% Renewable'].replace(r'[.]*', np.NaN, regex=True)
    Energy['Energy Supply per Capita'] = Energy['Energy Supply per Capita'].replace(r'[.]*', np.NaN, regex=True)
    #Corrections for country
    Energy['Country'] = Energy['Country'].str.replace(r' [(].*[)]',"")
    Energy['Country'] = Energy['Country'].str.replace(r"\d{1,3}","")
    Energy['Country']=Energy['Country'].replace({"Republic of Korea": "South Korea",
                                                 "United States of America": "United States",
                                                 "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                                                 "China, Hong Kong Special Administrative Region": "Hong Kong"})
    #Energy Done
    
    #GDP Start
    GDP = pd.read_csv('assets/world_bank.csv',
                      header=4)
    #Corrections for country
    GDP['Country Name'] = GDP['Country Name'].replace({"Korea, Rep.": "South Korea",
                                                       "Iran, Islamic Rep.": "Iran",
                                                       "Hong Kong SAR, China": "Hong Kong"})
    GDP.rename(columns={'Country Name':'Country'}, inplace=True)
    #GDP Done
    #ScimEn Start
    ScimEn = pd.read_excel('assets/scimagojr-3.xlsx')
    #ScimEn Done
    
    #Intersection Merges
    df = pd.merge(Energy, GDP, how='inner',on=['Country'])
    df_return = pd.merge(df,ScimEn,how='inner',on=['Country'])
    #Set Index
    df_return.set_index(['Country'], inplace=True)
    df_return = df_return[['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 
                           'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per Capita',
                           '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013',
                           '2014','2015']]
    #Sort
    df_return.sort_values(by='Rank', ascending=True, inplace=True)
    return df_return[:15]
    raise NotImplementedError()
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    df1 = Energy()
    df2 = GDP()
    df3 = ScimEn()
    
    EntriesOuter = len(pd.concat([df1['Country'],df2['Country'],df3['Country']]).unique())
    
    EntriesInner = (df1.merge(df2, left_on='Country', right_on='Country', how='inner')
                    .merge(df3, left_on='Country', right_on='Country', how='inner').shape[0])
            
    return EntriesOuter-EntriesInner
    raise NotImplementedError()
answer_two()
```

Output:
```
156
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    GDP = answer_one()
    GDP = GDP.iloc[:,10:]
    avgGDP = GDP.mean(axis=1)
    avgGDP.sort_values(axis=0, ascending=False,inplace=True)
    return avgGDP
    raise NotImplementedError()
answer_three()
```

Output:
```
Country
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
Iran                  4.441558e+11
dtype: float64
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    df = answer_three()
    key = df[5:6]
    key.index[0]
    GDP = pd.read_csv('assets/world_bank.csv',
                      header=4)
    GDP = GDP [['Country Name', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    GDP_Change = GDP.loc[GDP['Country Name'] == key.index[0]]
    Change = GDP_Change['2015'] - GDP_Change['2006']

    return Change.iloc[0]
    raise NotImplementedError()
answer_four()
```

Output:
```
246702696075.3999
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    df = answer_one()
    #Energy Done
    return float(df.iloc[:,8].mean())
    raise NotImplementedError()
answer_five()
```

Output:
```
157.6
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    df = answer_one()
    CountryRenew = df['% Renewable'].idxmax()
    MaxRenew = df['% Renewable'].max()
    
    TupCP = (CountryRenew, MaxRenew)
    return TupCP
    
    raise NotImplementedError()
answer_six()
```

Output:
```
('Brazil', 69.64803)
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    df = answer_one()
    df['Ratio'] = df['Self-citations']/df['Citations']
    Max = df['Ratio'].idxmax()
    Max_Tup = (Max, df['Ratio'].max())
    return Max_Tup
    raise NotImplementedError()
answer_seven()
```

Output:
```
('China', 0.6893126179389422)
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    df8 = Energy()
    df8['Population'] = df8['Energy Supply'] / df8['Energy Supply per Capita']
    df8.sort_values(by=['Population'],axis=0, ascending=False,inplace=True,)
    Third = df8[2:3]
    Third_Series = Third.iloc[0]
    return Third_Series[0]
    
    raise NotImplementedError()
answer_eight()
```

Output:
```
'United States'
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15 = Top15[['Energy Supply per Capita', 'Citable docs per Capita']]
    
    Corr = Top15.corr(method='pearson')
    
    return Corr.loc['Citable docs per Capita'].values[0]
    raise NotImplementedError()
answer_nine()
```

Output:
```
0.7940010435442943
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

```python
#plot9()
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    HighRenew = answer_one()
    HighRenew = HighRenew['% Renewable']
    HighRenew['Renewable?'] = HighRenew.apply(lambda x: 1 if x >= HighRenew.median() else 0)
    HighRenew = HighRenew['Renewable?']
    return HighRenew.rename('HighRenew')
    
    raise NotImplementedError()
list(answer_ten())
```

Output:
```
[1, 0, 0, 0, 1, 1, 1, 0, 1, 0, 1, 1, 0, 0, 1]
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

```python
ContinentDict  = {'China':'Asia', 
              'United States':'North America', 
              'Japan':'Asia', 
              'United Kingdom':'Europe', 
              'Russian Federation':'Europe', 
              'Canada':'North America', 
              'Germany':'Europe', 
              'India':'Asia',
              'France':'Europe', 
              'South Korea':'Asia', 
              'Italy':'Europe', 
              'Spain':'Europe', 
              'Iran':'Asia',
              'Australia':'Australia', 
              'Brazil':'South America'}
```

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    df11 = Energy()
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    df11['Population'] = df11['Energy Supply'] / df11['Energy Supply per Capita']
    df11.sort_values(by=['Population'],axis=0, ascending=False,inplace=True)
    df11['Continent'] = df11['Country'].map(ContinentDict)
    del df11['Energy Supply']
    del df11['Energy Supply per Capita']
    del df11['% Renewable']
    
    df11_Pop = df11.groupby(by=['Continent']).count()
    df11_Pop['sum'] = df11.groupby(by=['Continent']).sum()
    df11_Pop['mean'] = df11.groupby(by=['Continent']).mean()
    df11_Pop['std'] = df11.groupby(by=['Continent']).std()
    del df11_Pop['Population']
    df11_Pop.rename(columns={'Country':'size'}, inplace=True)
    return df11_Pop
    return mean
    raise NotImplementedError()
answer_eleven()
```

Output:
```
size           sum          mean           std
Continent                                                    
Asia              5  2.898666e+09  5.797333e+08  6.790979e+08
Australia         1  2.331602e+07  2.331602e+07           NaN
Europe            6  4.579297e+08  7.632161e+07  3.464767e+07
North America     2  3.528552e+08  1.764276e+08  1.996696e+08
South America     1  2.059153e+08  2.059153e+08           NaN
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    Top15 = answer_one()

    Top15['Continent'] = Top15.index.to_series().map(ContinentDict)
    Top15['% Renewable'] = pd.cut(Top15['% Renewable'],5)

    Top15.groupby(['Continent','% Renewable']).size().apply(lambda x: np.nan if x==0 else x).dropna()
    
    A12 = Top15.groupby(['Continent','% Renewable']).size()
    #A12.rename_axis({'Bins','% Renewable'}, inplace=True) #.apply(lambda x: np.nan if x==0 else x).dropna()
    return A12

    raise NotImplementedError()
answer_twelve()
```

Output:
```
Continent      % Renewable     
Asia           (2.212, 15.753]     4
               (15.753, 29.227]    1
Australia      (2.212, 15.753]     1
Europe         (2.212, 15.753]     1
               (15.753, 29.227]    3
               (29.227, 42.701]    2
North America  (2.212, 15.753]     1
               (56.174, 69.648]    1
South America  (56.174, 69.648]    1
dtype: int64
```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    Q13Top15 = answer_one()
    Q13Top15['PopEst'] = Q13Top15['Energy Supply'] / Q13Top15['Energy Supply per Capita']
    return Q13Top15['PopEst'].apply(lambda x: '{0:,}'.format(x)).astype(str)
    raise NotImplementedError()
answer_thirteen()
```

Output:
```
Country
China                 1,367,645,161.2903225
United States          317,615,384.61538464
Japan                  127,409,395.97315437
United Kingdom         63,870,967.741935484
Russian Federation            143,500,000.0
Canada                  35,239,864.86486486
Germany                 80,369,696.96969697
India                 1,276,730,769.2307692
France                  63,837,349.39759036
South Korea            49,805,429.864253394
Italy                  59,908,256.880733944
Spain                    46,443,396.2264151
Iran                    77,075,630.25210084
Australia              23,316,017.316017315
Brazil                 205,915,254.23728815
Name: PopEst, dtype: object
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```

```python
#plot_optional()
```

```python

```

```python

```